# COMP SCI 1015 IAP - W04 - Practical

# Task 1 - Price range

Assume we you are purchasing a laptop for the school. The school has decided the `lower` and `upper` limits of the purchase. Please write a function `purchase` that takes three parameters `lower`, `upper` and `laptop_cost` and decide wheather the proposed laptop is `too expensive`, `too cheap`, or `acceptable`. 

The function should return one of the following strings: `too cheap`, `too expensive`, or `acceptable`.



```python
def purchase(lower, upper, laptop_cost):
    # INSERT YOUR CODE BELOW
    # ~ 6 lines
    ...
        
s = purchase(500, 1000, 800) # return 'acceptable'
print(s)
s = purchase(1000, 1500, 2000) # return 'too expensive'
print(s)
```

# Task 2 - Estimate Pi

The mathmetician has found that the following infinite series can be used to generate a numerical approximation of $1/\pi$

$\frac{1}{\pi}=\frac{2\sqrt{2}}{9801}\sum_{k=0}^{\infty}\frac{(4k)!(1103+26390k)}{(k!)^{4}396^{4k}}$

Please complete the implementation of the `estimate_pi()` function below:

```python
import math

def estimate_pi(n):
    '''print the estimated pi value '''
    s = 0
    k = 0
    # INSERT YOUR CODE BELOW
    # ~ 6 lines
    while k <= n:
    ...
    # END OF YOUR CODE
    return 1/s

est = estimate_pi(0) # est_pi: 3.141592730013306 pi: 3.141592653589793     
print(f'est_pi: {est:.15f} pi: {math.pi:.15f}')
est = estimate_pi(3) # est_pi: 3.3.141592653589793 pi: 3.141592653589793     
print(f'est_pi: {est:.15f} pi: {math.pi:.15f}')
```

# Task 3: the three dice game

Please write a simple three dice program that keep rolling three dices until these three dices have the same value. The program then report how many times it roll the dices. This activity tests what you have learned so far and might be challenging. Let TA know if you need help!

```python
import random

def roll():    
    '''return 1 if all dices have same values. return 0 otherwise.'''
    dice1 = random.randint(1, 6) 
    dice2 = random.randint(1, 6)
    dice3 = random.randint(1, 6)    
    print(f'roll: {dice1} {dice2} {dice3}')
    
    # YOUR CODE BELOW
    # ~ 4 LINES
    ...
    
def play():
    '''play the game, keep rolling until three dices have same value'''
    counter = 0

    # YOUR CODE BELOW
    # ~ 5 LINES
    ...
    print(f'The player rolled {counter} times!')
    return counter

play()
```

Tests for Task 3 (points: 40)
```
>>> play() == 47
The player rolled 32 times!
False
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task1 (points: 30)
```
>>> purchase(500, 1000, 800) == 'acceptable'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> purchase(1000, 1500, 2000) == 'too expensive'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> purchase(800, 1000, 700) == 'too cheap'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
