```python
# Initialize Otter
import otter
grader = otter.Notebook("22S1_W05_Workshop.ipynb")
```

# COMP SCI 1015 IAP - W05 - Workshop

# Demo 1 - String  

A **string** consists of a series of **characters**. So far we already know that we can create strings by using the single quote `' '` or double quotes`" "`. We can also access each character with the bracket operators `[ ]` as shown below. 

```python
s = 'apple'
print(s[0]) # we call index starts from 0 
print(s[1])
print(s[2])
print(s[3])
print(s[4])
```

We call the expression in brackets an **index**. It is important to remember that the **index** starts from `0` in most popular programming language. That is, to access the first character of a string `s`, you should use `s[0]` instead of `s[1]`. 

## Length of a string and the last character 

`len` is a built-in function that returns the number of characters of a string.

```python
s = 'apple'
print(len(s))
```

Now, how can we access the **last** character of a string? You might be tempted to do the following

```python
l = len(s) # first get the length of the string
s[l] # access the last character
```

However, doing so will produce an index error. It is because the index starts from **zero**. Thus, to get the last character of the string, you should do

```python
l = len(s)
s[l-1]
```

## Travese the string 

Many applications process strings one character at a time. To visit every single character in a string, we can use the `while` loop from last week:

```python
i = 0
uni_name = 'University of Adelaide'

while i < len(uni_name):
    print(uni_name[i])
    i = i + 1
```

The while loop prints out each characters in a separate line. `i` is a variable storing the index value. In each iteration, the index `i` is increased by 1. The loop conditions on that the index value `i` is smaller than the length of the string. 

<!-- BEGIN QUESTION -->

# Activity 1 - Reverse the string 

Please complete the function below that prints out the reverse of the string. For example, if the argument is `apple`, your program should print out

e
l
p
p
a

```python
def reverse_str(s):
    # INSERT YOUR CODE BELOW
    # ~ 4lines
    ...

# TEST CASES BELOW
reverse_str('apple')
print('---')
reverse_str('happy')
```

<!-- END QUESTION -->

# Demo 2 - List

Like *string*, a **list** is a seuqence of values. However, the value in a *list* can be integers, float, string, or other Python data types. We call these values **elements** or **items**. 

The simplest way to create a list is to put the elements in square brackets `[ ]`. Code below contrast *string* and *list* in Python. 

```python
# STRING #

my_str = "apple" # values enclosed in double quotes
my_str[0] # access character with brackets


# LIST #

my_list = [2, 4, 6, 8, 10] # element enclosed in bracket
my_list[0] # access element with brackets, same as string
my_list[len(my_list)-1] # access element with brackets, same as string
```

List is a very powerful data structure and can contain multiple data types or even another list. However, to simplify your code, you might want to ensure your list contains only one data type. 

```python
my_list_2 = ['this', 'is', 'my', 'second', 'list!'] # element can be any data type

my_list_3 = ['hello', 100, [1, 3.5]] # this list contains string, integer, list!

print(my_list_2)
print(my_list_3)

```

## Travese a list

Similar to the traversal of a *string*, we can use a while loop to go over each element in a list.

```python
my_list = [2, 4, 6, 8, 10]
i = 0
while i < len(my_list):
    print(my_list[i])
    i = i + 1
```

<!-- BEGIN QUESTION -->

# Activity 2 - Calculating sum and mean of a list

Please write two functions that calculate the sum of every elements and its mean value.

```python
def my_sum(li):
    s = 0
    i = 0
    # COMPLETE THE CODE BELOW
    # ~ 3 LINES
    ...
    # END OF YOUR CODE
    return s
    
def my_mean(li):
    # COMPLETE THE CODE BELOW
    # You should use the my_sum function above
    # ~1 LINE
    m = ...
    return m
    
# TEST CASES BELOW  
li = [2, 4, 6, 8, 10]
print(my_sum(li)) # 30
print(my_mean(li)) # 6

li = [1, 3, 5, 7, 9, 12, 102, 12]
print(my_sum(li)) # 151
print(my_mean(li)) # 18.875
    
```

<!-- END QUESTION -->

***

# Demo 3 - List operations

## Update elements

Lists are *mutable*, meaning We can access and update the elements inside a list.

```python
l = [1, 2, 3, 4, 5]
print(l)
l[0] = 100 # update the first element
print(l)
```

Note that string is **immutable**. Although we can access character using the same bracket operator `[ ]`, but we <font color='red'>cannot</font> update the characters in a string. 

```python
s = "apple"
s[0] = "b"
```

## Update lists 

We can use the `+` and `*` operators on a list. The `+` operator concat lists while the `*` operator repeats a list multiple times.

```python
l1 = [1, 2, 3]
l2 = [4, 5, 6]
l3 = l1 + l2 # concat l1 and l2
print(l3)

l4 = l2 + l1 # concat l2 in front of l1
print(l4) 

l5 = l1 * 3 # repeat l1 three times.
print(l5)
```

## Remove elements 

The list provides a `pop` function and a `remove` function for removing the elements. The `pop(i)` remove `i-th` element while `remove(x)` remove an element with the value `x`

```python
l1 = [1, 2, 3]
l1.pop(0) #remove the first element
print(l1)

l2 = [2, 4, 6, 8]
l2.pop(3) # remove the 4th element (remember, the index starting from 0)
print(l2)

l3 = [2, 4, 6, 8]
l3.remove(6)
print(l3)
```

The `pop` function will generate errors if you try to pop an element out of the index range and the `remove` function will generate errors when you try to remove a non-existent value. 

## List slice 

The **slice** operator allows us to take a segement of a sequence. This operator works on both *list* and *string*. The syntax is `[n:m]` and returns the part of the sequence from the `n-th` value to `m-th` value, including the first but excluding the last.

```python
l1 = [2, 4, 6, 8, 10]
print(l1[0:3]) # elements 0, 1, 2
print(l1[2:5]) # elements 2, 3, 4

print('----')

# you can omit the first or second 
print(l1[:2]) # from the begining to 2nd element
print(l1[3:]) # from the 3rd element to the last element
```

Similarly, the slice operator also applies to a string

```python
s = "Adelaide"
print(s[3:6])
print(s[0:6])
print(s[3:])
```

<!-- BEGIN QUESTION -->

# Activity 3 - A classroom attendance recorder 

Let's create a system that records student entering and leaving the room using the list data structure. Please use a list that record the stduetns in the room. **Print the list** after each step.

1. Initially, Aaren, Charlie, and Ray are already in the room
1. 'Lane' come into the room
1. 'Jess' and 'Jude' come into the room
1. 'Lane' leaves the room
1. Finally, show the length of the list and its content

```python
#step 1
class_room = ...
print(class_room)
#step 2
class_room = ...
print(class_room)
#step 3
class_room = ...
print(class_room)
#step 4
class_room.remove('Lane') # solution
print(class_room)
#step 5
print(f'The number of remaining student is {...}')
print(f'The students in the class are {...}')

```

<!-- END QUESTION -->

***

# Demo 4 - For loop 

Earlier we have used a `while` loop to traverse both *string* and *list*. Another common way to traverse the sequence is via the `for` loop. Below I wrote two traversal approaches side-by-side. As you will see, the `for` loop provides a more succint way to iterate through the sequence.

```python
my_list = ['an', 'apple', 'a', 'day']

# USING WHILE 
print('--- while loop ---')
i = 0
while i < len(my_list):
    print(my_list[i])
    i = i + 1
    
print('--- for loop ---')
# USING FOR
for e in my_list:
    print(e)
    
```

## Range function 

However, there is a downside of the `for` loop - you do not have access to the **index** value, which could be crucial in some applications. For example, you might want to compare with the element after you etc. As an alternative, we also commonly use the `range` function with the `for` loop. 

The range() function generates a series of integer numbers between a **start** integer and a **stop** integer. 

```python
for i in range(5):
    print(i)
    
print('----')
for i in range(2, 5):
    print(i)
```

And we can utilise the values as the **index** to access a sequence.

```python
print('--- for with range function ---')
# FOR WITH RANGE
for i in range(len(my_list)): 
    print(my_list[i]) # use i as an index to access the values in the sequence.
```

Again, I put all three loops in a single cell for your reference.

```python
# USING WHILE 
print('--- while loop ---')
i = 0
while i < len(my_list):
    print(my_list[i])
    i = i + 1
    
print('--- for loop ---')
# USING FOR
for e in my_list:
    print(e)
    
print('--- for loop with range function ---')
# FOR WITH RANGE
for i in range(len(my_list)): 
    print(my_list[i]) # use i as an index to access the values in the sequence.
```

<!-- BEGIN QUESTION -->

# Activity 4 - Calculating sum and mean of a list

Please use `for` loop to complete the `my_sum` function.

```python
def my_sum(li):
    s = 0
    # COMPLETE THE CODE BELOW
    # ~ 2 LINES
    ...
    return s
    
def my_mean(li):
    m = ...
    return m
    
li = [2, 4, 6, 8, 10]
print(my_sum(li)) # 30
print(my_mean(li)) # 6

li = [1, 3, 5, 7, 9, 12, 102, 12]
print(my_sum(li)) # 151
print(my_mean(li)) # 18.875
    
```

<!-- END QUESTION -->

# Workshop Quiz: find largest number in a list

Please complete the function `find_max`, which returns the largest number in the given sequence. 

**hint:** `-math.inf` will give you an infinitly small number

Note, <font color='red'>using built-in max() function will not be rewarded any mark</font>

```python
import math


def find_max(seq):
    #INSERT YOUR CODE BELOW
    # ~4 LINES
    my_max = ...
    ...
    return my_max

#================
#TEST CASES BELOW
#================

li = [5, 6, 7, 8]
v = find_max(li)
print(f'answer": {max(li)} your answer: {v}')

li = [12, 34, 56, 98]
v = find_max(li)
print(f'answer": {max(li)} your answer: {v}')

li = [-321, -2933, -452, -1029]
v = find_max(li)
print(f'answer": {max(li)} your answer: {v}')
```

Tests for Activity (points: 100)
```
>>> find_max([3, 4, 5, 9]) == 9
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> find_max([-100, 100, 50]) == 100
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
