```python
try:
    from utils import *
except:
    pass
```

# COMP SCI 1015 IAP - W01 - Workshop submission

This is your first workshop submission. Please write the code that print out `Hello World!`

```python
# INPUT YOUR CODE BELOW
# 1 LINE
...
```

Output:
```
Hello World!
```

Tests for Task1 (points: 100)
```
>>> out = get_print_output(); # retrieves the print output.
>>> reset_print_output(); # Resets the print output.
>>> 'Hello World' in out
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
